<?xml version="1.0" encoding="UTF-8"?>
<quiz xmlns="http://stack-assessment.org/2025/moodle-question">>
  <question type="stack">
    <name>
      <text>Conversion of Complex Number: Cartesian to Polar and Exponential Forms</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the complex number \(z = {@aa@} {@sigb@} {@ev(abs(bb),simp)@}\cdot i\).</p>
<p>1. Write \(z\) in polar or trigonometric form. Give the modulus in exact form and the argument correct to \(2\) decimal places within your answer.</p>
<p>\(z = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>
<p>2. Write \(z\) in exponential form. Give the modulus and the argument in exact form. Recall that \(e^k\) is entered as <code>%e^(k)</code>.</p>
<p>\(z = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution:</h3>
<p>1. To write \(z\) in polar form, we need to find its modulus and argument.</p>
<p>\(\displaystyle r = |z| = \sqrt{{@lba@}{@aa@}{@rba@}^2 + {@lbb@}{@bb@}{@rbb@}^2}\)</p>
<p>\(\displaystyle  = \sqrt{{@aa^2+bb^2@}}\)</p>
<p>\(\displaystyle  = {@mod@}\).</p>
<p>To find the argument, we find the argument of \(\displaystyle {@ev(abs(aa),simp)@} + {@ev(abs(bb),simp)@}\cdot i\), which is</p>
<p>\(\displaystyle \tan^{-1} \left( \frac{{@ev(abs(bb),simp)@}}{{@ev(abs(aa),simp)@}} \right) = {@dispdp(primarg,2)@}\)</p>
<p>If not in the first quadrant, we need to use the symmetry in the Argand Diagram to find the actual argument, which happens to be \(\displaystyle {@argdp@}\) .</p>
<p>Therefore, the polar form of \(z\) is \( \displaystyle {@ta_1@}\) or \( \displaystyle {@ta1@}\)</p>

<p>2. To write \(z\) in exponential form, we need to find its modulus and argument.</p>
<p>From the previous part we know that \(\displaystyle r = |z| = {@mod@}\) and \(\theta = {@argdp@}\).</p>
<p>Therefore, the exponential form of \(z\) is \(\displaystyle {@ta2@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[simp:false;
pma:rand([-1,1]);
aa:ev(pma*(rand(6)+2),simp);
pmb:rand([-1,1]);
bb:ev(pmb*(rand(6)+2),simp);
sigb:if bb<0 then "-" else "+";
zz:aa+bb*i;
quad:if aa>0 then (if bb>0 then 1 else 4) else (if bb>0 then 2 else 3); 
primarg:ev(atan(abs(bb/aa)),simp);
arg:if quad = 1 then primarg elseif quad=2 then ev(%pi-primarg,simp) elseif quad=3 then ev(-(%pi-primarg),simp) else ev(-primarg,simp);
argdp:(ev(round(float(arg*100))/100,simp));
mod:ev(sqrt(aa^2+bb^2),simp);
/*worked solution variables*/
lba:if aa<0 then"(" else "";
rba:if aa<0 then ")" else "";
lbb:if bb<0 then "(" else "";
rbb:if bb<0 then ")" else "";
simp:true;
/*model answers*/
ta1: mod*(decimalplaces(cos(argdp),2) + i*decimalplaces(sin(argdp),2)); /*in 2dp*/
ta_1: mod*(cos(argdp) + i*sin(argdp)); 
ta2: mod*%e^(argdp*i);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(z\) in polar or trigonometric form \(={@ta1@}\)</p>
<p>\(z\) in exponential form \(={@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrectness</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>aa+i*bb</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is not in polar form.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks if the answer is correct but it's not given in the correct form</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mod*(cos(arg)+i*sin(arg))</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but it's not given in the correct form. Make sure you rounded the argument appropriately and you kept the modulus in exact form in terms of square roots.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>0</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>EqualComAss</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>sqrt(ev(aa^2+bb^2),simp)*%e^(argdp*i)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks dor incorrect answer</description>
        <answertest>EqualComAss</answertest>
        <sans>ans2</sans>
        <tans>aa+i*bb</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is not in exponential form.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for correct answer given in incorrect form</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but it's not given in the correct form. Make sure you rounded the argument appropriately and you kept the modulus in exact form in terms of square roots.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for correct answer given in incorrect form</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>aa+bb*i</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but it's not given in the correct form. Make sure you rounded the argument appropriately and you kept the modulus in exact form in terms of square roots.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>359828973</deployedseed>
    <deployedseed>943045251</deployedseed>
    <deployedseed>2020464580</deployedseed>
    <deployedseed>1446334543</deployedseed>
    <deployedseed>1261371823</deployedseed>
    <deployedseed>1589094950</deployedseed>
    <deployedseed>886953651</deployedseed>
    <deployedseed>1855696538</deployedseed>
    <deployedseed>955113234</deployedseed>
    <deployedseed>945443638</deployedseed>
    <deployedseed>2111441444</deployedseed>
    <deployedseed>1325169854</deployedseed>
    <deployedseed>315870026</deployedseed>
    <deployedseed>99016672</deployedseed>
    <deployedseed>1694284035</deployedseed>
    <deployedseed>1473051497</deployedseed>
    <deployedseed>118461861</deployedseed>
    <deployedseed>328648262</deployedseed>
    <deployedseed>1985607071</deployedseed>
    <deployedseed>1160828564</deployedseed>
    <deployedseed>363593940</deployedseed>
    <deployedseed>1552423245</deployedseed>
    <deployedseed>1940831511</deployedseed>
    <deployedseed>1264465677</deployedseed>
    <deployedseed>8632260</deployedseed>
    <deployedseed>1411296592</deployedseed>
    <deployedseed>1653750477</deployedseed>
    <deployedseed>937866726</deployedseed>
    <deployedseed>1420736145</deployedseed>
    <deployedseed>1117716935</deployedseed>
    <deployedseed>1565684087</deployedseed>
    <deployedseed>455533152</deployedseed>
    <deployedseed>324805589</deployedseed>
    <deployedseed>756012045</deployedseed>
    <deployedseed>216531957</deployedseed>
    <deployedseed>1886435243</deployedseed>
    <deployedseed>503685534</deployedseed>
    <deployedseed>1114898693</deployedseed>
    <deployedseed>1313241450</deployedseed>
    <deployedseed>24639811</deployedseed>
    <deployedseed>1496869959</deployedseed>
    <deployedseed>1999878899</deployedseed>
    <deployedseed>1399917375</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
