<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Factorization of polynomials reducible to quadratic ones</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the expression \({@expr@}\) and factorize it fully.</p>
<p>\({@expr@}=\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>\({@expr@}={@disp_coef_x2*xx^2@}+ ({@var_zz@})+({@var_yy@})+({@disp_coef_x0@})\)</p>
<p>\(\qquad \qquad \qquad \qquad\qquad={@c_21*xx@}({@c_11*xx-c_12@})+{@-c_22@} ({@c_11*xx-c_12@})\)</p>
<p>\(\qquad \qquad \qquad \qquad\qquad={@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/*Generating quadratic*/
c_11:rand([2,3,5]);
c_12:rand_with_prohib(-6,6,[0,c_11,2*c_11,3*c_11,-c_11,-2*c_11,-3*c_11]);
c_21:rand_with_prohib(2,5,[c_11,4]);
c_22:rand_with_prohib(-6,6,[0,c_21,2*c_21,3*c_21,-c_21,-2*c_21,-3*c_21]);
expr_fac:(c_11*x-c_12)*(c_21*x-c_22);
expr_exp:expand(expr_fac);
/*Hiding the quadratic by adding a^2b^4*/
xx:a^2*b^4;
expr:ev(expr_exp,x=xx);
/*Model answer*/
ta:factor(expr);

/*General Solution Variables*/
/*Finding coefficients - method 1*/
disp_coef_x2:coeff(expr_exp,x,2);
disp_coef_x1:coeff(expr_exp,x,1);
disp_coef_x0:coeff(expr_exp,x,0);

var_yy:-1*(c_11*c_22*xx);
var_zz:-1*(c_12*c_21*xx);
var_yz:var_zz+var_yy;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<br>\({@expr@}={@ta@} \)]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Emmaculate Odhiambo</p>
<p id="concept">[https://docs.google.com/document/d/1cakr9ozU1kZ6uCIEHViZyRc-1cWiku6g84cnsDssmuk/edit#heading=h.vy8x1tbmoha9]</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Factorizing quadratic polynomials of different variables, e.g. x, x^3, ab^2</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1249426572</deployedseed>
    <deployedseed>649850175</deployedseed>
    <deployedseed>161873925</deployedseed>
    <deployedseed>363641866</deployedseed>
    <deployedseed>521113067</deployedseed>
    <deployedseed>1135160175</deployedseed>
    <deployedseed>1768496100</deployedseed>
    <deployedseed>260433025</deployedseed>
    <deployedseed>1750240723</deployedseed>
    <deployedseed>1703485048</deployedseed>
    <deployedseed>1724939497</deployedseed>
    <deployedseed>1818743670</deployedseed>
    <deployedseed>946002737</deployedseed>
    <deployedseed>2132372359</deployedseed>
    <deployedseed>1727193477</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
