<?xml version="1.0" encoding="UTF-8"?>
<quiz xmlns="http://stack-assessment.org/2025/moodle-question">>
  <question type="stack">
    <name>
      <text>Finding the modulus of a complex number</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the modulus of the complex number \(\displaystyle z = \frac{{@num@}}{{@den@}}\)</p>
<p><em>Note: Give your answers in exact form i.e using fractions or writing in square root form</em></p>
<p>\(\displaystyle   |z| = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the modulus of a complex number, you can use the formula:</p>
 
<p>\(|z| = \sqrt{{\text{Re}(z)}^2 + {\text{Im}(z)}^2}\),</p>

<p>where \(\text{Re}(z)\) is the real part of the complex number and \(\text{Im}(z)\) is the imaginary part.</p>

<p>Given the complex number \(\displaystyle z = \frac{{@num@}}{{@den@}}\), let's simplify it first:</p>

<p>\(\displaystyle z = \frac{{@num@}}{{@den@}}\)</p>

<p>To get rid of the denominator, you can multiply both the numerator and the denominator by \({@den_conj@}\):</p>

<p>\(\displaystyle z = \left( \frac{{@num@}}{{@den@}} \right) \left( \frac{{@den_conj@}}{{@den_conj@}} \right) \)</p>
<p>\(\displaystyle z = \frac{{@expand(num1)@}}{{@den1@}}\)</p>
<p>\(\displaystyle z = {@expand(z_conjugate)@}\)</p>

<p>Now we have the real and imaginary parts of the complex number:</p>

<p>\( \displaystyle \text{Re}(z) = {@re_z@}\)</p>
<p>\(\displaystyle \text{Im}(z) = {@im_z@}\)</p>

<p>Now we can plug these values into the modulus formula:</p>

<p>\(\displaystyle \left | z \right |= \sqrt{( {@re_z@})^2 + \left({@im_z@}\right)^2}\)</p>
<p>\(\displaystyle \left | z \right | = \sqrt{ {@re_z^2@} + {@im_z^2@}}\)</p>
<p>\(\displaystyle |z| = \sqrt{ {@re_z^2 + im_z^2@}}\)</p>
<p>\(\displaystyle |z| = {@mod@}\)</p>

<p>So, the modulus of the complex number \(\displaystyle z = \ \frac{{@num@}}{{@den@}}\) is \( \displaystyle {@mod@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2023060500</text>
    </stackversion>
    <questionvariables>
      <text>/* Define the complex number */
a:rand(50)+1;
b:rand_with_prohib(-50,50,[0,a]);
num:(a + b*%i);
den:(-b*%i);
den_conj:conjugate(den);
z: num/den;


/* Simplify the complex number */
z_simplified: ratsimp(z);

/* Multiply by conjugate to get rid of the denominator */
num1:num*den_conj;
den1:den*den_conj;
z_conjugate:num1/den1;

/* Simplify the conjugate expression */
z_conjugate_simplified: (num1/den1);

/* Separate real and imaginary parts */
re_z: realpart(z_conjugate_simplified);
im_z: imagpart(z_conjugate_simplified);

/* Calculate the modulus */
mod: sqrt(re_z^2 + im_z^2);

/* Print the result */
ta:mod;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Find the modulus of the complex number \(z ={@z@}\).</p>
<p>\(\displaystyle  |z| = {@ta@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>715658396</deployedseed>
    <deployedseed>1796841613</deployedseed>
    <deployedseed>861923084</deployedseed>
    <deployedseed>1354833240</deployedseed>
    <deployedseed>2113675039</deployedseed>
    <deployedseed>441464700</deployedseed>
    <deployedseed>2110747280</deployedseed>
    <deployedseed>8138437</deployedseed>
    <deployedseed>113910359</deployedseed>
    <deployedseed>299019219</deployedseed>
    <deployedseed>949769970</deployedseed>
    <deployedseed>433795093</deployedseed>
    <deployedseed>2033362044</deployedseed>
    <deployedseed>1826061242</deployedseed>
    <deployedseed>1312659927</deployedseed>
    <deployedseed>1019756386</deployedseed>
    <deployedseed>1395889749</deployedseed>
    <deployedseed>452908940</deployedseed>
    <deployedseed>1206717253</deployedseed>
    <deployedseed>1316353304</deployedseed>
    <deployedseed>1179570601</deployedseed>
    <deployedseed>364289024</deployedseed>
    <deployedseed>352856768</deployedseed>
    <deployedseed>299010921</deployedseed>
    <deployedseed>993276230</deployedseed>
    <deployedseed>1386091915</deployedseed>
    <deployedseed>1614895386</deployedseed>
    <deployedseed>182924938</deployedseed>
    <deployedseed>16590934</deployedseed>
    <deployedseed>1766078073</deployedseed>
    <deployedseed>1384713634</deployedseed>
    <deployedseed>540856146</deployedseed>
    <deployedseed>1548417533</deployedseed>
    <deployedseed>28129705</deployedseed>
    <deployedseed>1715481665</deployedseed>
    <deployedseed>1886197550</deployedseed>
    <deployedseed>1105168639</deployedseed>
    <deployedseed>1241039430</deployedseed>
    <deployedseed>1844068760</deployedseed>
    <deployedseed>167261775</deployedseed>
    <deployedseed>94604449</deployedseed>
    <deployedseed>2092867510</deployedseed>
    <deployedseed>1535368722</deployedseed>
    <deployedseed>256174563</deployedseed>
    <deployedseed>622266808</deployedseed>
    <deployedseed>389609490</deployedseed>
    <deployedseed>68405089</deployedseed>
    <deployedseed>967584975</deployedseed>
    <deployedseed>944723990</deployedseed>
    <deployedseed>2143000489</deployedseed>
    <deployedseed>1412418563</deployedseed>
    <deployedseed>898429308</deployedseed>
    <deployedseed>1694951098</deployedseed>
    <deployedseed>278757446</deployedseed>
    <deployedseed>1162370384</deployedseed>
    <deployedseed>1726974609</deployedseed>
    <deployedseed>940366243</deployedseed>
    <deployedseed>2123008668</deployedseed>
    <deployedseed>457540339</deployedseed>
    <deployedseed>2085854277</deployedseed>
    <deployedseed>942866637</deployedseed>
    <deployedseed>838587968</deployedseed>
    <deployedseed>810100534</deployedseed>
    <deployedseed>1675072941</deployedseed>
    <deployedseed>1421188418</deployedseed>
    <deployedseed>1631643591</deployedseed>
    <deployedseed>979047499</deployedseed>
    <deployedseed>1040901498</deployedseed>
    <deployedseed>1771585727</deployedseed>
    <deployedseed>432901931</deployedseed>
    <deployedseed>629508097</deployedseed>
    <deployedseed>1135141234</deployedseed>
    <deployedseed>1741363864</deployedseed>
    <deployedseed>778643592</deployedseed>
    <deployedseed>1192584009</deployedseed>
    <deployedseed>1915825196</deployedseed>
    <deployedseed>1485858371</deployedseed>
    <deployedseed>431072562</deployedseed>
    <deployedseed>448281192</deployedseed>
    <deployedseed>1767875819</deployedseed>
    <deployedseed>815054825</deployedseed>
    <deployedseed>1245365402</deployedseed>
    <deployedseed>1517428749</deployedseed>
    <deployedseed>259698858</deployedseed>
    <deployedseed>2139550200</deployedseed>
    <deployedseed>1887368133</deployedseed>
    <deployedseed>917617600</deployedseed>
    <deployedseed>1482581681</deployedseed>
    <deployedseed>1721633840</deployedseed>
    <deployedseed>797475967</deployedseed>
    <deployedseed>1996985767</deployedseed>
    <deployedseed>988009355</deployedseed>
    <deployedseed>1302006280</deployedseed>
    <deployedseed>385314934</deployedseed>
    <deployedseed>898409936</deployedseed>
    <deployedseed>1530898774</deployedseed>
    <deployedseed>473076191</deployedseed>
    <deployedseed>1395017319</deployedseed>
    <deployedseed>2115977603</deployedseed>
    <deployedseed>1046686953</deployedseed>
    <deployedseed>493409041</deployedseed>
    <deployedseed>548040811</deployedseed>
    <deployedseed>1475930816</deployedseed>
    <deployedseed>180237215</deployedseed>
    <deployedseed>876573890</deployedseed>
    <deployedseed>1455282268</deployedseed>
    <deployedseed>1000879728</deployedseed>
    <deployedseed>1507367037</deployedseed>
    <deployedseed>1433702155</deployedseed>
    <deployedseed>73028070</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
