<?xml version="1.0" encoding="UTF-8"?>
<quiz xmlns="http://stack-assessment.org/2025/moodle-question">>
  <question type="stack">
    <name>
      <text>Evaluation of Logarithmic Expression without Calculator</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find, without using a calculator, the value of \( \log_{{@aa@}}{@cc@}\).</p>
<p>Type the intermediary steps in your answer to show your working. The syntax for logarithms is as follows: \( \log_{a} b\) is entered as logbase(b,a). Please seek advice or post in the issues forum if you have a problem with this notation.</p>
<p>\( \log_{{@aa@}}{@cc@}=\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked solution</h3>
<p>\(\begin{align} \log_{{@aa@}}{@cc@} &= \log_{{@aa@}} {@aa@}^{{@bb@}}\\
& = {@bb@}\cdot \log_{{@aa@}} {@aa@} \\
&= {@bb@} \cdot 1 \\
&= {@bb@}\\
\end{align}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>aa:rand([2,3,4,5,6]);
bb:rand([2,3]);
cc:ev(aa^bb,simp);
ta:[lg(aa^bb,aa),stackeq(bb*lg(aa,aa)),stackeq(bb)];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>\(\begin{align} \log_{{@aa@}}{@cc@} &= \log_{{@aa@}} {@aa@}^{{@bb@}}\\
& = {@bb@}\cdot \log_{{@aa@}} {@aa@} \\
&= {@bb@} \cdot 1 \\
&= {@bb@}\\
\end{align}\).</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Emmaculate</p>
<p id="concept">[URL of Concept Doc]</p>
<p id="reviewer">Godfrey Wabwire</p>
<p id="reviewer">Wastalas Motognon</p>
<p id="description">Given an logarithmic expression, need to evaluate without using a calculator</p>]]></text>
    </questiondescription>
    <questionsimplify>0</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><i class="fa fa-check"></i> Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><i class="fa fa-adjust"></i> Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><i class="fa fa-times"></i> Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>equiv</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>5</insertstars>
      <syntaxhint>[lg(?),=?*lg(?),=?]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Comparing if sa is equal to logarithmic expression</description>
        <answertest>EqualComAss</answertest>
        <sans>last(ans1)</sans>
        <tans>last(ta)</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct evaluation of logarithmic expression</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1[2]</sans>
        <tans>ta[2]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for the correct answer</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1[2]</sans>
        <tans>ta[3]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You are not showing enough working.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>checks for the correct final answer with missing  intermediary steps</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>stackeq(bb)</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer is correct but you seem to be missing intermediary steps in your workings!!</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <deployedseed>6</deployedseed>
    <deployedseed>7</deployedseed>
    <deployedseed>8</deployedseed>
    <deployedseed>9</deployedseed>
    <deployedseed>10</deployedseed>
    <deployedseed>11</deployedseed>
    <deployedseed>12</deployedseed>
    <deployedseed>13</deployedseed>
    <deployedseed>14</deployedseed>
    <deployedseed>15</deployedseed>
    <deployedseed>16</deployedseed>
    <deployedseed>17</deployedseed>
    <deployedseed>18</deployedseed>
    <deployedseed>19</deployedseed>
    <deployedseed>20</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
